# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    listx = list(x)
    listy = list(y)
    countx = {
        '1': 0,
        '2': 0,
        '3': 0,
        '4': 0,
        '5': 0,
        '6': 0,
        '7': 0,
        '8': 0,
        '9': 0       
    }
    
    county = {
        '1': 0,
        '2': 0,
        '3': 0,
        '4': 0,
        '5': 0,
        '6': 0,
        '7': 0,
        '8': 0,
        '9': 0
    }
    for i in list(x):
        if i in countx:
            countx[i] += 1
    for i in list(y):
        if i in county:
            county[i] += 1
    if countx == county:
        return True
    else:
        return False
same_digits('232','322')
```

Output:
```
True
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
#Find the smallest positive integer, x, such that 2x, 3x, 4x, 5x, and 6x, contain the same digits.

i = 1
i2 = i*2
i3 = i*3
i4 = i*4
i5 = i*5
i6 = i*6
for i in range(0,1000000):
    if same_digits(str(i),str(i2)) == True:
        if same_digits(str(i2),str(i3)) == True:
            if same_digits(str(i3),str(i4)) == True:
                if same_digits(str(i4),str(i5)) == True:
                    if same_digits(str(i5),str(i6)) == True:
                        print(i)
                        break
    else:
        i +=1
        i2 = i*2
        i3 = i*3
        i4 = i*4
        i5 = i*5
        i6 = i*6
#THESE LINES ARE JUST TO MANUALLY CHECK THE NUMBERS TO MAKE SURE IT WORKS
# print(i)
# print(i2)
# print(i3)
# print(i4)
# print(i5)
# print(i6)

#ANSWER is 142857

```

Output:
```
142857
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```