# Algorithms Exercise 1

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import numpy as np
```

## Word counting

Write a function `tokenize` that takes a string of English text returns a list of words. It should also remove [stop words](http://en.wikipedia.org/wiki/Stop_words), which are common short words that are often removed before natural language processing. Your function should have the following logic:

* Split the string into lines using `splitlines`.
* Split each line into a list of words and merge the lists for each line.
* Use Python's builtin `filter` function to remove all punctuation.
* If `stop_words` is a list, remove all occurences of the words in the list.
* If `stop_words` is a space delimeted string of words, split them and remove them.
* Remove any remaining empty words.
* Make all words lowercase.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def makelist(s):
    lines = s.splitlines()
    lst = []
    for i in lines:
        for j in i.split():
            lst.append(j)
    return lst
def fil(func,lst):
    newlist = []
    for i in lst:
        newlist.append(list(filter(func,i)))
    string = ''
    newthing=[]
    for i in newlist:
        for j in i:
            string += j
        newthing.append(string)
        string = ''
    return newthing
def inlist(lst,words):
    thing = []
    for i in lst:
        if i not in words:
            thing.append(i)
    return thing

def makelower(lst):
    thing = []
    for i in lst:
        thing.append(i.lower())
    return thing

def tokenize(s, stop_words=None, punctuation='`~!@#$%^&*()_-+={[}]|\:;"<,>.?/}\t'):
    """Split a string into a list of words, removing punctuation and stop words."""
        
    def punc(i):
        if i in punctuation:
            return False
        else:
            return True  
    print(fil(punc,makelist(test)))  
    lst = makelist(s)
    thing = fil(punc,lst)
    thing = makelower(thing)
    if isinstance(stop_words,list) == True:
        thing = inlist(thing,stop_words)
    if isinstance(stop_words,str) == True:
        stuff = makelist(stop_words)
        thing = inlist(thing,stuff)
    return thing

test = 'wor!d1\n word2--word3 \nword4'
def punc(i):
    if i in '`~!@#$%^&*()_-+={[}]|\:;"<,>.?/}\t':
        return False
    else:
        return True
print(makelist(test))

newlst = []
for i in makelist(test):
    for k in i:
        if k == '-':
            l = i.split('-')            
for k in l:
    newlst.append(k)
thing = [i for i in newlst if i != '']
listed = []
for i in makelist(test):
    for k in i:
        if k == '-':
            for m in thing:
                listed.append(m)
            break 
    listed.append(i)
print(listed)
```

Output:
```
['wor!d1', 'word2--word3', 'word4']
['wor!d1', 'word2', 'word3', 'word2--word3', 'word4']
```

<!-- grade_id: algorithmsex01a, points: 2, grade: true, solution: false -->
```python
assert tokenize("This, is the way; that things will end", stop_words=['the', 'is']) == \
    ['this', 'way', 'that', 'things', 'will', 'end']
wasteland = """
APRIL is the cruellest month, breeding
Lilacs out of the dead land, mixing
Memory and desire, stirring
Dull roots with spring rain.
"""
assert tokenize(wasteland, stop_words='is the of and') == \
    ['april','cruellest','month','breeding','lilacs','out','dead','land',
     'mixing','memory','desire','stirring','dull','roots','with','spring',
     'rain']
```

Output:
```
['word1', 'word2word3', 'word4']
['word1', 'word2word3', 'word4']
```


Write a function `count_words` that takes a list of words and returns a dictionary where the keys in the dictionary are the unique words in the list and the values are the word counts.

<!-- grade_id: algorithmsex01b, points: 2, grade: false, solution: true -->
```python
def count_words(data):
    """Return a word count dictionary from the list of words in data."""
    words = []
    count = []
    for i in data:
        words.append(i)
        count.append(data.count(i)) 
    return dict(zip(words,count))
```

<!-- grade_id: algorithmsex01b, points: 2, grade: true, solution: false -->
```python
assert count_words(tokenize('this and the this from and a a a')) == \
    {'a': 3, 'and': 2, 'from': 1, 'the': 1, 'this': 2}
```

Output:
```
['word1', 'word2word3', 'word4']
```

Write a function `sort_word_counts` that return a list of sorted word counts:

* Each element of the list should be a `(word, count)` tuple.
* The list should be sorted by the word counts, with the higest counts coming first.
* To perform this sort, look at using the `sorted` function with a custom `key` and `reverse`
  argument.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def sort_word_counts(wc):
    """Return a list of 2-tuples of (word, count), sorted by count descending."""
    def isValue(b):
        return b[1]
    a = list(wc.items())
    return sorted(a,key=isValue,reverse=True)
```

<!-- grade_id: algorithmsex01c, points: 2, grade: true, solution: false -->
```python
assert sort_word_counts(count_words(tokenize('this and a the this this and a a a'))) == \
    [('a', 4), ('this', 3), ('and', 2), ('the', 1)]
```

Output:
```
['word1', 'word2word3', 'word4']
```

Perform a word count analysis on Chapter 1 of Moby Dick, whose text can be found in the file `mobydick_chapter1.txt`:

* Read the file into a string.
* Tokenize with stop words of `'the of and a to in is it that as'`.
* Perform a word count, the sort and save the result in a variable named `swc`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
f = open("mobydick_chapter1.txt",'r')
string = ''
for i in f:
    string += i
words = tokenize(string,stop_words='the of and a to in is it that as')
wordcount = count_words(words)
swc = sort_word_counts(wordcount)
```

Output:
```
['word1', 'word2word3', 'word4']
```

<!-- grade_id: algorithmsex01d, points: 2, grade: true, solution: false -->
```python
assert swc[0]==('i',43)
assert len(swc)==848
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-669-79f30673d9f4> in <module>()
      1 assert swc[0]==('i',43)
----> 2 assert len(swc)==848

AssertionError:
```

Create a "Cleveland Style" [dotplot](http://en.wikipedia.org/wiki/Dot_plot_%28statistics%29) of the counts of the top 50 words using Matplotlib. If you don't know what a dotplot is, you will have to do some research...

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: algorithsex01e, points: 2, grade: true, solution: false -->
```python
assert True # use this for grading the dotplot
```