# Numpy Exercise 2

## Imports

```python
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
import seaborn as sns
```

## Factorial

Write a function that computes the factorial of small numbers using `np.arange` and `np.cumprod`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def np_fact(n):
    if n == 0:
        return 1
    num = np.arange(1,n+1)
    fact = num.cumprod()
    return list(fact).pop()
#check
[np_fact(i) for i in range(0,11)]
#good
```

Output:
```
[1, 1, 2, 6, 24, 120, 720, 5040, 40320, 362880, 3628800]
```

<!-- grade_id: numpyex02a, points: 4, grade: true, solution: false -->
```python
assert np_fact(0)==1
assert np_fact(1)==1
assert np_fact(10)==3628800
assert [np_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Write a function that computes the factorial of small numbers using a Python loop.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def loop_fact(n):
    if n == 0:
        return 1
    total = 1
    for i in range(1,n+1):
        total = total*i
    return total
loop_fact(10)
```

Output:
```
3628800
```

<!-- grade_id: numpyex02b, points: 4, grade: true, solution: false -->
```python
assert loop_fact(0)==1
assert loop_fact(1)==1
assert loop_fact(10)==3628800
assert [loop_fact(i) for i in range(0,11)]==[1,1,2,6,24,120,720,5040,40320,362880,3628800]
```

Use the `%timeit` magic to time both versions of this function for an argument of `50`. The syntax for `%timeit` is:

```python
%timeit -n1 -r1 function_to_time()
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
print("Using regular old python loops I got")
%timeit -n1 -r1 loop_fact(50)
print("Using the new fancy NumPy stuff I got")
%timeit -n1 -r1 np_fact(50)
```

Output:
```
Using regular old python loops I got
1 loops, best of 1: 16 µs per loop
Using the new fancy NumPy stuff I got
1 loops, best of 1: 65.1 µs per loop
```

In the cell below, summarize your timing tests. Which version is faster? Why do you think that version is faster?

<!-- grade_id: numpyex02c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

It appears that the NumPy calculation takes longer than the regular old python loop stuff. In calculating the NumPy factorial I had to use np.cumprod, which appears to generate a list of cumulative products, not just the last value (which is why i had to list().pop() it out). These extra calculations likely slow down the overall calculation.