# Project Euler: Problem 17

https://projecteuler.net/problem=17

If the numbers 1 to 5 are written out in words: one, two, three, four, five, then there are 3 + 3 + 5 + 4 + 4 = 19 letters used in total.

If all the numbers from 1 to 1000 (one thousand) inclusive were written out in words, how many letters would be used?


NOTE: Do not count spaces or hyphens. For example, 342 (three hundred and forty-two) contains 23 letters and 115 (one hundred and fifteen) contains 20 letters. The use of "and" when writing out numbers is in compliance with British usage.

First write a `number_to_words(n)` function that takes an integer `n` between 1 and 1000 inclusive and returns a list of words for the number as described above

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def number_to_words(n):
    """Given a number n between 1-1000 inclusive return a list of words for the number."""
    """Got idea to use dictionary from gitter"""
    lst = []
    dic20 = {
        1: 'one',
        2: 'two',
        3: 'three',
        4: 'four',
        5: 'five',
        6: 'six',
        7: 'seven',
        8: 'eight',
        9: 'nine',
        10: 'ten',
        11: 'eleven',
        12: 'twelve',
        13: 'thirteen',
        14: 'fourteen',
        15: 'fifteen',
        16: 'sixteen',
        17: 'seventeen',
        18: 'eighteen',
        19: 'nineteen',
    }
    dic100 = {
        20: 'twenty',
        30: 'thirty',
        40: 'forty',
        50: 'fifty',
        60: 'sixty',
        70: 'seventy',
        80: 'eighty',
        90: 'ninety',
    }
    dic1000 = {
        100: 'one hundred',
        200: 'two hundred',
        300: 'three hundred',
        400: 'four hundred',
        500: 'five hundred',
        600: 'six hundred',
        700: 'seven hundred',
        800: 'eight hundred',
        900: 'nine hundred',
    }
    if n < 20:
        for i in range(1,n+1):
            lst.append(dic20[i])
    elif n < 100:
        for i in range(1,20):
            lst.append(dic20[i])
        for j in range(20,n+1,10):
            for k in range(0,10):
                if j + k > n:
                    break
                if k == 0:
                    lst.append(dic100[j])
                else:
                    lst.append(dic100[j] + '-' + dic20[k])
    elif n <= 1000:
        add_1000 = False
        if n == 1000:
            add_1000 = True
            n = 999
        for i in range(1,20):
            lst.append(dic20[i])
        for j in range(20,100,10):
            for k in range(0,10):
                if k == 0:
                    lst.append(dic100[j])
                else:
                    lst.append(dic100[j] + '-' + dic20[k])
        for l in range(100,n+1,100):
            if n <120:
                for i in range(0,20):
                    if i + l > n:
                        break
                    if i == 0:
                        lst.append(dic1000[100])
                    else:
                        lst.append(dic1000[100] + " and " + dic20[i])
            else:
                for i in range(0,20):
                    if i + l > n:
                        break
                    if i == 0:
                        lst.append(dic1000[l])
                    else:
                        lst.append(dic1000[l] + " and " + dic20[i])
                for j in range(20,100,10):
                    for k in range(0,10):
                        if l + j + k > n:
                            break
                        if k == 0:
                            lst.append(dic1000[l] + " and " + dic100[j])
                        else:
                            lst.append(dic1000[l] + " and " + dic100[j] + '-' + dic20[k])
        if add_1000:
            lst.append('one thousand')
                
    return lst
len(number_to_words(1000))
```

Output:
```
1000
```

Now write a set of `assert` tests for your `number_to_words` function that verifies that it is working as expected.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
for i in range(1,999):
    assert len(number_to_words(i))==i;
```

```python
number_to_words(999)[-1]
```

Output:
```
'nine hundred and ninety-nine'
```

<!-- grade_id: projecteuler17a, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the number_to_words tests.
```

Now define a `count_letters(n)` that returns the number of letters used to write out the words for all of the the numbers `1` to `n` inclusive.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def count_letters(n):
    count = 0
    """Count the number of letters used to write out the words for 1-n inclusive."""
    for i in number_to_words(n)[:]:
        for j in i:
            if j != " " and j != "-":
                count += 1
    return count
print(count_letters(1000))
```

Output:
```
21124
```

Now write a set of `assert` tests for your `count_letters` function that verifies that it is working as expected.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
assert count_letters(5)==19
assert count_letters(1000)==21124
```

<!-- grade_id: projecteuler17b, points: 4, grade: true, solution: false -->
```python
assert True # use this for grading the count_letters tests.
```

Finally used your `count_letters` function to solve the original question.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
print(count_letters(5))
```

Output:
```
19
```

<!-- grade_id: projecteuler17c, points: 2, grade: true, solution: false -->
```python
assert True # use this for gradig the answer to the original question.
```